---
title: 'LEARNING PROGRAMMING WITH PYTHON'
description: "The resulting C code was then processed by a C compiler to produce an executable program. (C++\ncompilers today translate C++ directly into machine language.)"
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
